Before you turn this lab in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

**Provide your name and any collaborators below:**

<!-- grade_id: name-collabs, points: 0, grade: true, solution: true -->
```scala
// YOUR CODE HERE
// Zhou Ye
```

# Lab 5 - Inheritence
> Labs will be due each week before the homework. They are not intended take a significant amount of time but rather to provide examples/practice on specific and isolated features in the language. Labs are autograded so you can get quick feedback.

### Import the necessary Chisel dependencies. 
> There will be a cell like this in every lab. Make sure you run it before proceeding to bring the Chisel Library into the Jupyter Notebook scope!

```scala
interp.load.module(os.Path(s"${System.getProperty("user.dir")}/resource/chisel_deps.sc"))
```

```scala
import chisel3._
import chisel3.util._
import chiseltest._
import chiseltest.RawTester.test
```

Output:
```
[32mimport [39m[36mchisel3._
[39m
[32mimport [39m[36mchisel3.util._
[39m
[32mimport [39m[36mchiseltest._
[39m
[32mimport [39m[36mchiseltest.RawTester.test[39m
```

## Problem 1 (6 pts) - Method Templated for Chisel Types
> Write a generic `connect` method that takes in two generic arguments `a` and `b` of type `chisel3.Data` and connects `a` (input) to `b` (output).

<!-- grade_id: prob1, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def connect[T <: Data](b: T, a: T): Unit = {
    b := a
}
```

Output:
```
defined [32mfunction[39m [36mconnect[39m
```

<!-- grade_id: prob1-test, points: 6, grade: true, solution: false -->
```scala
class TestConnect extends Module {
    val io = IO(new Bundle {
        val u = Output(UInt())
        val s = Output(SInt())
    })
     
    connect(io.u, 10.U)
    connect(io.s, -10.S)
}

def testConnect: Boolean = {
    test(new TestConnect) { dut =>
        dut.io.u.expect(10.U)
        dut.io.s.expect(-10.S)
    }
    true
}
assert(testConnect)
```

Output:
```
defined [32mclass[39m [36mTestConnect[39m
defined [32mfunction[39m [36mtestConnect[39m
```

## Problem 2 (8 pts) - Case Class pattern matching and traits
> Fill in the body of `ProcessRequest` using pattern matching to translate a requested `MealItem` into a more descriptive string. The meal items are case classes inheriting from MealItem. You can infer what the strings should be by looking at the following testcases.

<!-- grade_id: prob2, points: 0, grade: false, solution: true -->
```scala
abstract class MealItem

case class Drink(variety: String) extends MealItem

case class Burger(withCheese: Boolean) extends MealItem

case object Fries extends MealItem

def ProcessRequest(request: MealItem): String = {
    request match {
        // YOUR CODE HERE
        case Drink("water") => "Water is healthy"
        case Drink("soda")  => "Soda is refreshing"
        case Burger(true)   => "Mmm... cheeseburger"
        case Burger(false)  => "Mmm... plain hamburger"
        case Fries          => "Fries were a good decision"
        case _              => s"Couldn't match $request"
    }
}

def ProcessOrder(order: Seq[MealItem]) {
    order foreach { item => println(ProcessRequest(item)) }
}
```

Output:
```
defined [32mclass[39m [36mMealItem[39m
defined [32mclass[39m [36mDrink[39m
defined [32mclass[39m [36mBurger[39m
defined [32mobject[39m [36mFries[39m
defined [32mfunction[39m [36mProcessRequest[39m
defined [32mfunction[39m [36mProcessOrder[39m
```

<!-- grade_id: prob2-tests, points: 8, grade: true, solution: false -->
```scala
assert (ProcessRequest(Drink("water")) == "Water is healthy", "water")

assert (ProcessRequest(Drink("soda")) == "Soda is refreshing", "soda")

assert (ProcessRequest(Burger(true)) == "Mmm... cheeseburger", "cheeseburger")

assert (ProcessRequest(Burger(false)) == "Mmm... plain hamburger", "hamburger")

assert (ProcessRequest(Fries) == "Fries were a good decision", "fries")

assert (ProcessRequest(Drink("special request")) == "Couldn't match Drink(special request)", "other")
```

## Problem 3 (11 pts) - Chisel trait
> Fill in the missing fields in the `DbgCounter` trait. The purpose of this trait is to add a debug counter that prints out the current cycle if the `debug` input (Bool) is true. The trait should add an output `count` (UInt) representing the debug counter. _Hint:_ peek at the test case to infer signal names and data types.

<!-- grade_id: prob6, points: 0, grade: false, solution: true -->
```scala
trait DbgCounter extends Module {
    // YOUR CODE HERE
    val debug = IO(Input(Bool()))
    val count = IO(Output(UInt(8.W)))
    def n : Int
    val (cnt, wrap) = Counter(true.B, n)
    count := cnt
    when (debug) {
        printf("%d\n", count)
    }
}
```

Output:
```
defined [32mtrait[39m [36mDbgCounter[39m
```

<!-- grade_id: prob6-tests, points: 11, grade: true, solution: false -->
```scala
class TestMod(maxCycles: Int) extends Module with DbgCounter {
   def n = maxCycles 
}

def testDbgCounter(n: Int, printDbg: Boolean): Boolean = {
    test(new TestMod(n)) { dut =>
        dut.debug.poke(printDbg.B)
        for (i <- 0 until n) { 
            dut.count.expect(i.U)
            dut.clock.step()
        }
        // test that counter wraps
        dut.count.expect(0.U)
    }
    true
}

assert(testDbgCounter(n=5, true))
assert(testDbgCounter(n=15, false))
```

Output:
```
0
   1
   2
   3
   4
defined [32mclass[39m [36mTestMod[39m
defined [32mfunction[39m [36mtestDbgCounter[39m
```